```python
# Initialize Otter
import otter
grader = otter.Notebook("discussion.ipynb")
```

```python

%load_ext autoreload
%autoreload 2
```

# Discussion 4

### Due Saturday October 23, 11:59:59PM



---


```python
# import libraries
import pandas as pd
import numpy as np
import os
from IPython.display import HTML
```

```python
# for formatting purposes
def multi_table(table_list):
    ''' Acceps a list of IpyTable objects and returns a table which contains each IpyTable in a cell
    '''
    return HTML(
        '<table><tr style="background-color:white;">' + 
        ''.join(['<td>' + table._repr_html_() + '</td>' for table in table_list]) +
        '</tr></table>'
    )
```

```python
from discussion import *
```

# Review: Combining DataFrames 

#### `merge()`

* Used to combine two (or more) dataframes on the basis of **values of common columns** (indices can also be used, use `left_index=True` and/or `right_index=True`).
    * If we are joining columns on columns, the DataFrame indexes will be ignored. 
    * If we are joining indexes on indexes or indexes on a column or columns, the index will be passed on.

```python
# left dataframe
left = pd.DataFrame({
   'id':[1,2,3,4,5],
   'Name': ['Aaron', 'Marina', 'Justin', 'Janine', 'Ilya'],
   'subject_id':['sub1','sub2','sub4','sub6','sub5']})

# right dataframe
right = pd.DataFrame(
   {'id':[1,2,3,4,5],
   'Name': ['Enrique', 'Parker', 'Erik', 'Allston', 'Betty'],
   'subject_id':['sub2','sub4','sub3','sub6','sub5']})

multi_table([left, right])
```

Output:
```
<IPython.core.display.HTML object>
```

* **`on`**: column or index level names to join on. 
    * These must be found in both DataFrames. 
    * If `on` is `None` and not merging on indexes then this defaults to the intersection of the columns in both DataFrames.

```python
# merge left and right tables on 'id' column
on_id = pd.merge(left,right,on='id')

# how many rows, how many columns?
multi_table([left, right, on_id])
```

Output:
```
<IPython.core.display.HTML object>
```

```python
# merge left and right tables on 'id' and 'subject_id' columns
on_id_subject = pd.merge(left,right,on=['id', 'subject_id'])

# how many rows, how many columns, what are the indices?
multi_table([left, right, on_id_subject])
```

Output:
```
<IPython.core.display.HTML object>
```

* **`how`**: specifies how to determine which keys are to be included in the resulting table. 
    * If a key (column name) combination does not appear in either the left or the right tables, the values in the joined table will be `np.NaN`.
    * Defaults to `inner`, joining will be performed on index. 

```python
data_a = {
        'subject_id': ['1', '2', '3', '4', '5'],
        'first_name': ['Manny', 'Will', 'Hunter', 'Ian', 'Eric'], 
        'last_name': ['Machado', 'Myers', 'Renfroe', 'Kinsler', 'Hosmer']}
df_a = pd.DataFrame(data_a, columns = ['subject_id', 'first_name', 'last_name'])

data_b = {
        'subject_id': ['4', '5', '6', '7', '8'],
        'first_name': ['Cody', 'Justin', 'Corey', 'Clayton', 'Kenley'], 
        'last_name': ['Bellinger', 'Turner', 'Seager', 'Kershaw', 'Jansen']}
df_b = pd.DataFrame(data_b, columns = ['subject_id', 'first_name', 'last_name'])

multi_table([df_a, df_b])
```

Output:
```
<IPython.core.display.HTML object>
```

| Merge Method  | Description                  |
| :-------      | :---------------------------:| 
| `left`        | Use keys from left object    | 
| `right`       | Use keys from right object   | 
| `outer`       | Use union of keys            |
| `inner`       | Use intersection of keys     | 

```python
# based on the output below, what 'how' argument was passed into pd.merge?
how_list = ['outer', 'inner', 'right', 'left']

merge_method = np.random.choice(how_list)

pd.merge(df_a, df_b, on='subject_id', how=merge_method)
```

Output:
```
subject_id first_name_x last_name_x first_name_y last_name_y
0          1        Manny     Machado          NaN         NaN
1          2         Will       Myers          NaN         NaN
2          3       Hunter     Renfroe          NaN         NaN
3          4          Ian     Kinsler         Cody   Bellinger
4          5         Eric      Hosmer       Justin      Turner
```

```python
# let's check!
merge_method
```

Output:
```
'left'
```

#### `concat()`

* Used to append one (or more) dataframes one below the other (or sideways, depending on whether the axis option is set to 0 or 1).
    * Useful if we have two or more data sets containing the same columns but different rows of data.
    * We can also the columns from one `Dataframe` to those of another `Dataframe`.

```python
# add 'left' below 'right'
pd.concat([right, left])
```

Output:
```
id     Name subject_id
0   1  Enrique       sub2
1   2   Parker       sub4
2   3     Erik       sub3
3   4  Allston       sub6
4   5    Betty       sub5
0   1    Aaron       sub1
1   2   Marina       sub2
2   3   Justin       sub4
3   4   Janine       sub6
4   5     Ilya       sub5
```

```python
# if you want to keep track of the names dataframes after concat, use 'keys'
pd.concat([right, left], keys=['right', 'left'])
```

Output:
```
id     Name subject_id
right 0   1  Enrique       sub2
      1   2   Parker       sub4
      2   3     Erik       sub3
      3   4  Allston       sub6
      4   5    Betty       sub5
left  0   1    Aaron       sub1
      1   2   Marina       sub2
      2   3   Justin       sub4
      3   4   Janine       sub6
      4   5     Ilya       sub5
```

```python
# add 'left' to the side of 'right'
pd.concat([right, left], axis=1)
```

Output:
```
id     Name subject_id  id    Name subject_id
0   1  Enrique       sub2   1   Aaron       sub1
1   2   Parker       sub4   2  Marina       sub2
2   3     Erik       sub3   3  Justin       sub4
3   4  Allston       sub6   4  Janine       sub6
4   5    Betty       sub5   5    Ilya       sub5
```

#### `join()`

* Used to merge two dataframes on the basis of the index; instead of using `merge()` with the option `left_index=True` we can use `join()`.
    * Join operation honors the object on which it is called: `a.join(b)` $ \neq $ `b.join(a)`.

<img src="imgs/join_types.jpg">

1. **Inner Join** – default behavior, only keep rows where the merge “on” value exists in both the left and right dataframes.
2. **Left Outer** – keep every row in the left dataframe.
    * Where there are missing values of the “on” variable in the right dataframe, add `np.NaN` values in the result.
3. **Right Join** – keep every row in the right dataframe. 
    * Where there are missing values of the “on” variable in the left column, add `np.NaN` values in the result.
4. **Outer Join** – returns all the rows from the left dataframe, all the rows from the right dataframe, and matches up rows where possible, with `NaNs` elsewhere.

We'll start with a simple example:

```python
left = pd.DataFrame({'key': ['foo', 'bar'], 'val': [1, 2]}).set_index('key')
right = pd.DataFrame({'key': ['foo', 'bar'], 'val': [4, 5]}).set_index('key')

joined = left.join(right, lsuffix='_l', rsuffix='_r')

multi_table([left, right, joined])
```

Output:
```
<IPython.core.display.HTML object>
```

Now let's try something a bit more complex:

```python
df1_data = {
    'Year' : [2014, 2014, 2014, 2014, 2014],
    'Week' : ['A', 'B', 'B', 'C', 'D'],
    'Color' : ['Red', 'Red', 'Black', 'Red', 'Green'],
    'Val' : [50, 60, 70, 10, 20]
}

df1 = pd.DataFrame(df1_data).set_index('Week')

df2_data = {
    'Year' : [2014, 2014, 2014, 2014, 2014],
    'Week' : ['A', 'B', 'C', 'C', 'D'],
    'Color' : ['Black', 'Black', 'Green', 'Red', 'Red'],
    'Score' : [30, 100, 50, 20, 40]
}

df2 = pd.DataFrame(df2_data).set_index('Week')

multi_table([df1, df2])
```

Output:
```
<IPython.core.display.HTML object>
```

```python
# how many rows, how many columns?
df1.join(df2, lsuffix='_l', rsuffix = '_r')
```

Output:
```
Year_l Color_l  Val  Year_r Color_r  Score
Week                                            
A       2014     Red   50    2014   Black     30
B       2014     Red   60    2014   Black    100
B       2014   Black   70    2014   Black    100
C       2014     Red   10    2014   Green     50
C       2014     Red   10    2014     Red     20
D       2014   Green   20    2014     Red     40
```

```python
# will this be any different?
df2.join(df1, lsuffix='_l', rsuffix = '_r')
```

Output:
```
Year_l Color_l  Score  Year_r Color_r  Val
Week                                            
A       2014   Black     30    2014     Red   50
B       2014   Black    100    2014     Red   60
B       2014   Black    100    2014   Black   70
C       2014   Green     50    2014     Red   10
C       2014     Red     20    2014     Red   10
D       2014     Red     40    2014   Green   20
```

---

## Practice Problems

**Question 1**

You are given two seperate dataframes: `mlb_2017` and `mlb_2018`. Both dataframes contain statistics for the 2017 and 2018 baseball seasons respectively. Your job is two combine these two dataframes into one using the following guidelines:

* The dataframe you return should be indexed by team name (`Tm`).
* The dataframe you return should include all columns from both `mlb_2017` and `mlb_2018`.
* Use the suffixes `_2017` and `_2018` to differentiate between statistics from both seasons.

Create a function `combined_seasons` that returns, as a tuple, the following:

* The combined dataframe described above.
* The team with most homeruns (`HR`) bewteen the 2017 and 2018 seasons combined.

```python
# read in the following .txt files
mlb_2017 = pd.read_csv(os.path.join('data','mlb_2017.txt'))
mlb_2018 = pd.read_csv(os.path.join('data','mlb_2018.txt'))

multi_table([mlb_2017.head(), mlb_2018.head()])
```

Output:
```
<IPython.core.display.HTML object>
```

```python

```

```python
m = mlb_2017.merge(mlb_2018, on = 'Tm', suffixes =['_2017','_2018'])

```

```python
q1_out[0].columns
```

Output:
```
Index(['#Bat_2017', 'BatAge_2017', 'R/G_2017', 'G_2017', 'PA_2017', 'AB_2017',
       'R_2017', 'H_2017', '2B_2017', '3B_2017', 'HR_2017', 'RBI_2017',
       'SB_2017', 'CS_2017', 'BB_2017', 'SO_2017', 'BA_2017', 'OBP_2017',
       'SLG_2017', 'OPS_2017', 'OPS+_2017', 'TB_2017', 'GDP_2017', 'HBP_2017',
       'SH_2017', 'SF_2017', 'IBB_2017', 'LOB_2017', '#Bat_2018',
       'BatAge_2018', 'R/G_2018', 'G_2018', 'PA_2018', 'AB_2018', 'R_2018',
       'H_2018', '2B_2018', '3B_2018', 'HR_2018', 'RBI_2018', 'SB_2018',
       'CS_2018', 'BB_2018', 'SO_2018', 'BA_2018', 'OBP_2018', 'SLG_2018',
       'OPS_2018', 'OPS+_2018', 'TB_2018', 'GDP_2018', 'HBP_2018', 'SH_2018',
       'SF_2018', 'IBB_2018', 'LOB_2018'],
      dtype='object')
```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
mlb_2017 = pd.read_csv(os.path.join('data','mlb_2017.txt'))
mlb_2018 = pd.read_csv(os.path.join('data','mlb_2018.txt'))
q1_out = combined_seasons(mlb_2017, mlb_2018)
```

```python
q1_out[0].shape
```

Output:
```
(30, 56)
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2**

Using the same two dataframes, `mlb_2017` and `mlb_2018`, create a function `seasonal_average` that combines them and takes the mean of each column for each team. 

* The dataframe you return should be indexed by team name (`Tm`).
* Each column should contain the average value between the *2017* and *2018* seasons for the given statistic for each team.
    * For example, the `HR` column should contain the average value for `HR` for each team between the *2017* and *2018* seasons.

```python

```

Output:
```
#Bat  BatAge    R/G      G      PA      AB      R       H     2B    3B  \
Tm                                                                            
ARI  47.0   28.75  4.645  162.0  6190.5  5492.5  752.5  1344.0  286.5  44.5   
ATL  53.5   28.00  4.605  162.0  6233.5  5583.0  745.5  1450.0  301.5  27.5   
BAL  53.0   28.50  4.215  162.0  6087.0  5578.5  682.5  1393.0  255.5  13.5   
BOS  46.5   27.50  5.130  162.0  6320.0  5646.0  830.5  1485.0  328.5  25.0   
CHC  48.5   27.15  4.870  162.5  6326.0  5560.0  791.5  1427.5  280.0  31.5   
CHW  51.0   26.60  4.205  162.0  6065.0  5518.0  681.0  1372.0  257.5  38.5   
CIN  50.0   27.15  4.475  162.0  6226.5  5508.0  724.5  1397.0  250.0  31.5   
CLE  45.0   28.70  5.050  162.0  6267.0  5553.0  818.0  1448.0  315.0  24.0   
COL  41.0   28.50  4.940  162.5  6189.5  5537.5  802.0  1464.0  286.5  40.0   
DET  49.0   28.75  4.215  162.0  6089.5  5525.0  682.5  1380.5  286.5  35.0   
HOU  43.5   28.50  5.225  162.0  6208.5  5532.0  846.5  1
```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
mlb_2017 = pd.read_csv(os.path.join('data','mlb_2017.txt'))
mlb_2018 = pd.read_csv(os.path.join('data','mlb_2018.txt'))
q2_out = seasonal_average(mlb_2017, mlb_2018)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## Congratulations! You're done!

* Submit your `.py` file to Gradescope. Note that you only need to submit the `.py` file; this notebook should not be uploaded. Make sure that all of your work is in the `.py` file and not here by running the doctests: `python -m doctest discussion.py`.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!
```

```python

```

Tests for q1 (points: None)
```
>>> q1_out[0].shape == (30, 56)
True
```
failure_message: doctest: shape
```
>>> q1_out[1] in q1_out[0].index
True
```
failure_message: doctest: index
```
>>> all([(('_2017' in x) or ('_2018' in x)) for x in q1_out[0]])
True
```
failure_message: doctest: columns
```
>>> q1_out[1] == 'NYY'
True
```
failure_message: doctest: most homruns incorrect


Tests for q2 (points: None)
```
>>> q2_out.shape == (30, 28)
True
```
failure_message: doctest: shape
```
>>> q2_out.index.nunique() == 30
True
```
failure_message: doctest: index
```
>>> q2_out.loc['MIN']['HR'] == 186
True
```
failure_message: doctest: MIN HR
```
>>> q2_out.loc['OAK'].max() == 6190.5
True
```
failure_message: doctest: OAK row
