```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

# DSC 80 - Lab 02

### Due Date: Monday October 11th, 11:59 PM

## Instructions
Much like in DSC 10, this Jupyter Notebook contains the statements of the problems and provides code and markdown cells to display your answers to the problems. Unlike DSC 10, the notebook is *only* for displaying a readable version of your final answers. The coding will be done in an accompanying `lab.py` file that is imported into the current notebook.

Labs and programming assignments will be graded in (at most) two ways:
1. The functions and classes in the accompanying python file will be tested (a la DSC 20),
2. The notebook may be graded (if it contains free response questions or asks you to draw plots).

**Note**: Labs will have public tests and private tests. The public "smoke tests" that you will run below and which appear on Gradescope are generally worth no points. After the due date, we will replace these tests with private tests that will determine your grade. This is different from DSC 10, where labs only had public tests!

**Do not change the function names in the `*.py` file**
- The functions in the `*.py` file are how your assignment is graded, and they are graded by their name.
- If you changed something you weren't supposed to, just use git to revert! Ask us if you need help with this, or google around for `git revert`.

**Tips for working in the Notebook**:
- The notebooks serve to present the questions and give you a place to present your results for later review.
- The notebook on *lab assignments* are not graded (only the `.py` file).
- Notebooks for *projects* will serve as a final report for the assignment, and contain conclusions and answers to open ended questions that are graded.
- The notebook serves as a nice environment for 'pre-development' and experimentation before designing your function in your `.py` file. You can write code here, but make sure that all of your real work is in the .py file.

**Tips for developing in the .py file**:
- Do not change the function names in the starter code; grading is done using these function names.
- Do not change the docstrings in the functions. These are there to tell you if your work is on the right track!
- You are encouraged to write your own additional helper functions to solve the lab! 
    - Developing in python usually consists of larger files, with many short functions.
    - You may write your other functions in an additional `.py` file that you import in `lab.py` (much like we do in the notebook).
- Always document your code!

### Importing code from `lab.py`

* We import our `.py` file that's contained in the same directory as this notebook.
* We use the `autoreload` notebook extension to make changes to our `lab.py` file immediately available in our notebook. Without this extension, we would need to restart the notebook kernel to see any changes to `lab.py` in the notebook.
    - `autoreload` is necessary because, upon import, `lab.py` is compiled to bytecode (in the directory `__pycache__`). Subsequent imports of `lab` merely import the existing compiled python.

```python
%load_ext autoreload
%autoreload 2
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

```python
from lab import *
```

```python
import os
import pandas as pd
import numpy as np
```

## Pandas Basics

**Question 1: Test scores**

You will be given a small dataset (so that you can manually check the correctness of your code). Please follow a few requirements when solving the problems below:

* For all questions you need to write code general enough to be applied to another similar dataset. 
* Do not hard-code any answers please. 
* Do not use `for` or `while` loops
   

1. Write a function called `data_load` that takes a file name of the data set to be read as a string and returns a dataframe following the steps below:

    a. Read only a subset of columns: `name`, `tries`, `highest_score`, `sex`
    
    b. Then you realized that for your analysis the column `sex` is not needed. Remove it. 
    
    c. You want to customize the column names: rename `name` to `firstname` and `tries` to `attempts`
    
    d. Turn the `firstname` column into the index.


2. Write a function `pass_fail` that takes the dataframe returned from the function above and adds a column `pass` based on the following conditions:

    * "Yes" if a number of attempts is strictly less than 3 and the score is >= 50
    * "Yes" if a number of attempts is strictly less than 6 and the score is >= 70
    * "Yes" if a number of attempts is strictly less than 10 and the score is >= 90
    * "No" otherwise
 
Your function should return the (modified) input dataframe with the added column.
    
3. Write a fuction `av_score` that takes in a dataframe from the question above and returns the average score for those students who passed the test. 
    
4. Write a function `highest_score_name` that takes in the dataframe from question 1.2 and returns a dictionary, where the key is the highest score and the value is the name (as a list) of the person with the highest score (attempts do not count). If more than one student got the highest score, include all names in a list. 

5. Write a function `idx_dup` that does not take any parameters and returns a single integer, answering the question below:

Is it possible for a dataframe's index to have duplicate values?
1. No, the index values must be unique and uses non-negative integers only, just like in numpy arrays
2. No, the index values must be unique and uses integers only
3. No, the index values must be unique but index values are not restricted to integers
4. Yes, but index values must be non-negative integers only
5. Yes, but index values must be integers only
6. Yes and index values are not restricted to integers
    


```python

```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
scores_fp = os.path.join('data', 'scores.csv')
scores = data_load(scores_fp)
passfail = pass_fail(scores.copy())
avscore = av_score(passfail.copy())
highest = highest_score_name(passfail)
idxdup = idx_dup()
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## Tricky Pandas.

Sometimes you can get input that you do not expect. The next set of questions walk you through a few examples that might surprise you. 

**Question 2 : Duplicate and selection**



1. Write a function `trick_me` that does not take any parameters. <br>Inside the function: 
    * Create a dataframe `tricky_1` that has three columns labeled: "Name", "Name", "Age". Your table should have 5 rows, the values are up to you. 
    * Save this dataframe in the `csv` file called `tricky_1.csv` without the index. 
    * Now create another dataframe, `tricky_2`, by reading in the file `tricky_1.csv `. What are your observations?
        1. It was not possible to create a dataframe with the duplicate columns
        2. `tricky_1` and `tricky_2` have the same column names
        3. `tricky_1` and `tricky_2` have different column names
    * Return your answer as a letter
    
    

2. Write a function `reason_dup` that answers the following question: `Why does pandas allow us to have duplicate column names?` by returning a corresponding letter. 
    1. It does not, duplicate column names are not allowed
    2. Since duplicate indices are allowed and we also can transpose a dataframe.
    3. It is a bug in Pandas
    
    
   
   
3. Write a function `trick_bool` that does not take any parameters. To determine the correct answers from the list below, you should follow the steps outlined by experimenting in *the notebook* (or a python REPL). Outside the function:
    * Create a dataframe `bools` that has four columns labeled: "True", "True", "False", "False". Each column name is boolean.
    * Your table should have 4 rows, the values are up to you. 
    * You need to think (without running it) what output you should get when running each line of code below. Pick a corresponding answer from a given list. Your function should return a list with three letters that correspond to the dataframe structure for each line below. 
    
     ```
     df[True]
     df[[True, True, False, False]]
     df[[True, False]]
     ```
    
        1. Dataframe: 2 columns, 1 row
        2. Dataframe: 2 columns, 2 rows
        3. Dataframe: 2 columns, 3 rows
        4. Dataframe: 2 columns, 4 rows
        5. Dataframe: 3 columns, 1 rows
        6. Dataframe: 3 columns, 2 rows
        7. Dataframe: 3 columns, 3 rows
        8. Dataframe: 3 columns, 4 rows
        9. Dataframe: 4 columns, 1 rows
        10. Dataframe: 4 columns, 2 rows
        11. Dataframe: 4 columns, 3 rows
        12. Dataframe: 4 columns, 4 rows
        13. Error
    
    
4.  Write a function `reason_bool` that answers the following question: `Why the outputs are the way they are?` by returning a corresponding letter. 
    1. booleans arrays select either rows or columns, randomly
    2. booleans arrays always select rows by default
    3. booleans arrays always select columns by default 
    4. booleans arrays always select rows by default, unless column names are set to `True`/`False` values.
    
    
    
   


```python
data = [[1,2,3,4]]*4
df = pd.DataFrame(data, columns=[True, True, False, False])
df
```

Output:
```
True   True   False  False
0      1      2      3      4
1      1      2      3      4
2      1      2      3      4
3      1      2      3      4
```

```python
df[True]
```

Output:
```
True  True
0     1     2
1     1     2
2     1     2
3     1     2
```

```python
df[[False, True, False, False]]
```

Output:
```
True   True   False  False
1      1      2      3      4
```

```python
df[[True, False]]
```

Output:
```
E:\anaconda\lib\site-packages\pandas\core\frame.py in __getitem__(self, key)
   3013         # Do we have a (boolean) 1d indexer?
   3014         if com.is_bool_indexer(key):
-> 3015             return self._getitem_bool_array(key)
   3016 
   3017         # We are left with two options: a single key, and a collection of keys,

E:\anaconda\lib\site-packages\pandas\core\frame.py in _getitem_bool_array(self, key)
   3060             )
   3061         elif len(key) != len(self.index):
-> 3062             raise ValueError(
   3063                 f"Item wrong length {len(key)} instead of {len(self.index)}."
   3064             )

ValueError: Item wrong length 2 instead of 4.
```

```python
# don't change this cell -- it is needed for the tests to work
trick_ans = trick_bool()
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

**Question 3 : np.NaN in a dataframe**


In the notebook, use the code given below to create a dataframe called `nans`. Note that we use `np.NaN` (`numpy`'s representation of 'Not a Number') to create missing values.
 
```
nans = pd.DataFrame([[0,1,np.NaN], [np.NaN, np.NaN, np.NaN], [1, 2, 3]])
```
Now you decided to make your dataset more readable for people who do not understand `NaN` and replace it with a `MISSING` string instead. In order to do that you wrote the following function:

```
def change(x):
    if x == np.NaN:
        return "MISSING"
    else:
        return x
```

* Write a line of code that applies the function above to the last column of the `nans` dataframe. 
* What was a result?
    * A: It worked: all np.NaNs in the last columns where changed to "MISSING"
    * B: It did not work: does not matter how I tried, the NaN values were not changed.
    
I expect you to answer `B` here. What had happened? Turns out, you can't use simple comparison `==` to detect if a value is `np.NaN`. You need to use another way to compare a variable to a `np.NaN`, read about it [here](https://stackoverflow.com/questions/41342609/the-difference-between-comparison-to-np-nan-and-isnull)

1. Modify the function `change` above to work as expected.
2. Write method `correct_replacement` that takes in a dataframe like `nans` and returns a modified dataframe, where all the `NaN` are replaced with `"MISSING"`. Use your corrected version of `change` to do this. **The pandas function .fillna is not allowed in this question.** 


```python
nans = pd.DataFrame([[0,1,np.NaN], [np.NaN, np.NaN, np.NaN], [1, 2, 3]])
nans
```

Output:
```
0    1    2
0  0.0  1.0  NaN
1  NaN  NaN  NaN
2  1.0  2.0  3.0
```

```python
def change(x):
    if np.isnan(x):
        return "MISSING"
    else:
        return x
```

```python
correct_replacement(nans)
```

Output:
```
0        1        2
0      0.0      1.0  MISSING
1  MISSING  MISSING  MISSING
2      1.0      2.0      3.0
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

### Summary Statistics

**Question 4**

In this question you will create two general purpose functions that make it easy to 'qualitatively' assess the contents of a dataframe.

1. Create a function `population_stats` which takes in a dataframe `df` and returns a dataframe indexed by the columns of `df`, with the following columns:
    * `num_nonnull` contains the number of non-null entries in each column,
    * `pct_nonnull` contains the proportion of entries in each column that are non-null,
    * `num_distinct` contains the number of distinct entries in each column,
    * `pct_distinct` contains the proportion of (non-null) entries in each column that are distinct from each other.
    
*Note*: you may find the `.nunique()` series method useful.

*Note*: The number of distinct entries does not include nulls.
    
2. Create a function `most_common` which takes in a dataframe `df` and a number `N` and returns a dataframe of the `N` most-common values (and their counts) for each column of `df`. Any column with fewer than `N` distinct values should contain `NaN` in those entries. 

*Note*: you can loop through the *columns* of `df` to construct your output. You should **not** be looping through rows.

For example, for the subset of the `salaries` dataframe with columns 'Job Title' and 'Status' from lecture one (left), `most_common(salaries, N=5)` is given (right). 

<table><tr>
    <td><img src="data/imgs/dataframe.png" width="70%"/></td>
    <td><img src="data/imgs/most_common.png" width="70%"/></td>
</tr></table>

```python

```

```python

```

```python

```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
pop_data = np.random.choice(range(10), size=(100, 4))
df_pop = pd.DataFrame(pop_data, columns='A B C D'.split())
out_pop = population_stats(df_pop)

common_data = np.random.choice(range(10), size=(100, 2))
common_df = pd.DataFrame(common_data, columns='A B'.split())
common_out = most_common(common_df, N=3)
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

**Question 5**

## Faulty Scooters

A new electric scooter company 'Maxwell Scooters' opened a retail shop in La Jolla recently and 300 UCSD students bought new scooters for getting around campus. After 8 students start complaining their scooters are faulty, negative on-line reviews for the scooters start to spread. In response, the scooter company adamantly claims that 99% of their scooters come off the production line working properly. You think this seems unlikely and decide to investigate.

* Select a significance level for you investigation. (Not to be turned in)
* What are reasonable choices for the *Null Hypothesis* for your investigation? Select all that apply:
    1. The scooter company produces scooters that are 99% non-faulty.
    2. The scooter company produces scooters that are less than 99% non-faulty.
    3. The scooter company produces scooters that are at least 1% faulty.
    4. The scooter company produces scooters that are ~2.6% faulty.

Return your answer in a function `null_hypoth` that takes zero arguments.

* Create a function `simulate_null` simulates a single step of data generation under the null hypothesis. The function should return a binary array.

* Create a function `estimate_p_val` that takes in a number `N` and returns the estimated p-value of your investigation upon simulating the null hypothesis `N` times.

*Note*: Plot the Null distribution and your observed statistic to check your work.

```python

```

Output:
```
5
```

```python
estimate_p_val(1000)
```

Output:
```
0.993
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

# Super-Heroes

The questions below analyze a dataset of super-heroes found in the `data` directory. One of the datasets have a list of attributes on each super-hero, while the other is a *boolean* dataframe of which super-heroes have which super-powers. Note, the datasets contain information on both *good* super-heroes, as well as *bad* super-heroes (AKA villains). 

**Question 6**

### Super-hero powers

Now read in the dataset of super-hero powers in the `data` directory. Create a function `super_hero_powers` that takes in a dataframe like `powers` and returns a list with the following three entries:

1. The name of the super-hero with the greatest number of powers.
2. The name of the most common super-power among super-heroes whose names begin with 'M'.
3. The most popular super-power among those with only one super-power.

You should *not* be hard-coding your answers in this question; your function should work on any dataset similar to `powers`. You should not be using loops in this question.

*Note:* You may find the `.idxmax` method useful in this problem.

```python

```

```python
super_hero_powers(powers)
```

Output:
```
['563', 'Super Strength', 'Intelligence']
```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
super_fp = os.path.join('data', 'superheroes_powers.csv')
powers = pd.read_csv(super_fp)
super_out = super_hero_powers(powers)
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

**Question 7**

### Super-hero attributes

Read in the dataset of super-hero attributes from the `data` directory. Use your summary functions from question 4 to help acquaint yourself with the dataset.

Cleaning the data: the dataset has no explicit null (`np.NaN`) values, although many entries *should* be null. Replace these values with null by creating a function `clean_heroes`.

Now answer the following questions, collecting your answers in a (function `super_hero_stats` that returns) a list. You should answer the questions using the *cleaned* super-heroes data; your answers *should* be hard-coded in the function.
1. Which publisher has a greater proportion of 'bad' characters -- Marvel Comics or DC Comics?
2. Give the number of characters that are NOT human, or the publisher is not Marvel Comics nor DC comics. For this question, only consider race "Human" as human, races such as "Human / Radiation" don't count as human.
3. Give the name of the character that's both greater than one standard deviation above mean in height and at least one standard deviation below the mean in weight.
4. Who is heavier on average: good or bad characters?
5. What is the name of the tallest Mutant with no hair?
6. What is the probability that a randomly chosen 'Marvel' character in the dataset is a woman?

*Note:* Since your answers to these questions should be hard-coded, you should not include your code in your .py file. Just return a list with your answers.

*Note:* Nan denotes an unknown value that does not count as an entry with any attributes.

```python
M = clean_out[clean_out['Publisher'] == 'Marvel Comics']
D = clean_out[clean_out['Publisher'] == 'DC Comics']
M_B = len(M[M['Alignment']=='bad'])/len(M)
D_B = len(D[D['Alignment']=='bad'])/len(D)
[M_B, D_B]
```

Output:
```
[0.2963917525773196, 0.2744186046511628]
```

```python
clean_out[((clean_out['Publisher'] != 'Marvel Comics') & (clean_out['Publisher'] != 'DC Comics' ))| (clean_out['Race'] != 'Human')]
```

Output:
```
name Gender Eye color               Race Hair color  Height  \
1     Abe Sapien   Male      blue      Icthyo Sapien    No Hair   191.0   
2       Abin Sur   Male      blue            Ungaran    No Hair   185.0   
3    Abomination   Male     green  Human / Radiation    No Hair   203.0   
4        Abraxas   Male      blue      Cosmic Entity      Black     NaN   
6    Adam Monroe   Male      blue                NaN      Blond     NaN   
..           ...    ...       ...                ...        ...     ...   
726        X-Man   Male      blue                NaN      Brown   175.0   
727  Yellow Claw   Male      blue                NaN    No Hair   188.0   
730         Ymir   Male     white        Frost Giant    No Hair   304.8   
731         Yoda   Male     brown     Yoda's species      White    66.0   
733         Zoom   Male       red                NaN      Brown   185.0   

             Publisher Skin color Alignment  Weight  
1    Dark Horse Comics       blue      good    65.0  
2  
```

```python
height_sd = np.std(clean_out['Height'])
height_mean = np.mean(clean_out['Height'])
weight_sd = np.std(clean_out['Weight'])
weight_mean = np.mean(clean_out['Weight'])
clean_out[(clean_out['Height'] >= (height_mean + height_sd)) & (clean_out['Weight'] <= (weight_mean - weight_sd)) ]
```

Output:
```
name Gender Eye color            Race Hair color  Height      Publisher  \
302  Groot   Male    yellow  Flora Colossus        NaN   701.0  Marvel Comics   

    Skin color Alignment  Weight  
302        NaN      good     4.0
```

```python
np.mean(clean_out[clean_out['Alignment']=='bad']['Weight'])
```

Output:
```
139.80985915492957
```

```python
np.mean(clean_out[clean_out['Alignment']=='good']['Weight'])
```

Output:
```
95.54654654654655
```

```python
clean_out[clean_out['Race']=='Mutant'].sort_values(by = ['Height'], ascending=False)
```

Output:
```
name  Gender Eye color    Race Hair color  Height  \
504                  Onslaught    Male       red  Mutant    No Hair   305.0   
184                   Colossus    Male    silver  Mutant      Black   226.0   
709                    Warpath    Male     brown  Mutant      Black   218.0   
34                  Apocalypse    Male       red  Mutant      Black   213.0   
144                      Cable    Male      blue  Mutant      White   203.0   
..                         ...     ...       ...     ...        ...     ...   
450               Michelangelo    Male      blue  Mutant        NaN     NaN   
488  Negasonic Teenage Warhead  Female     black  Mutant      Black     NaN   
526                  Power Man    Male       NaN  Mutant        NaN     NaN   
543                    Raphael    Male       NaN  Mutant    No Hair     NaN   
584             Sebastian Shaw    Male       NaN  Mutant        NaN     NaN   

          Publisher Skin color Alignment  Weight  
504   Marvel Comics       
```

```python
len(M[M['Gender'] == 'Female'])/len(M)
```

Output:
```
0.2860824742268041
```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
superheroes_fp = os.path.join('data', 'superheroes.csv')
heroes = pd.read_csv(superheroes_fp, index_col=0)
clean_out = clean_heroes(heroes)

stats_out = super_hero_stats()
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

**Question 8**

### Are blond-haired, blue-eyed characters disproportionately 'good'?

1. Create a function `bhbe` ('blond-hair-blue-eyes') that returns a boolean column that labels super-heroes/villains that are blond-haired *and* blue eyed.
    * Look at the values of the hair/eyes columns; it needs some cleaning! (The doctest makes sure you've cleaned it properly. If you don't pass the doctest, look more closely at the values in the columns!)


Now, you'd like to answer the question 
> "Are blond-haired, blue-eyed characters disproportionately 'good'?"

To do this, you'd like to test the null hypothesis:
> "The proportion of 'good' heroes among blond-haired, blue-eyed heroes is roughly the same as (equals) the proportion of 'good' heroes in the overall population."

Fix a significance level of 1%.

2. Create a function `observed_stat` that takes in `heroes`, and returns the observed test statistic.
3. Create a function `simulate_bhbe_null` that takes in a number `n` that returns a `n` instances of the test statistic generated under the null hypothesis. You should hard-code your simulation parameter into the function (rounding to the nearest hundredth is fine); the function should *not* read in any data.
4. Create a function `calc_pval` that returns a list where:
    * the first element is the p-value for hypothesis test (using 100,000 simulations). Please run the code yourself and hard-code this answer, as actually running the 100,000 simulation hypothesis test will timeout on gradescope. 
    * the second element is `Reject` if you reject the null hypothesis and `Fail to reject` if you fail to reject the null hypothesis.

```python
superheroes_fp = os.path.join('data', 'superheroes.csv')
heroes = pd.read_csv(superheroes_fp, index_col=0)
pd.set_option("display.max_rows", 999)
heroes
```

Output:
```
name  Gender                Eye color  \
0                       A-Bomb    Male                   yellow   
1                   Abe Sapien    Male                     blue   
2                     Abin Sur    Male                     blue   
3                  Abomination    Male                    green   
4                      Abraxas    Male                     blue   
5                Absorbing Man    Male                     blue   
6                  Adam Monroe    Male                     blue   
7                 Adam Strange    Male                     blue   
8                     Agent 13  Female                     blue   
9                    Agent Bob    Male                    brown   
10                  Agent Zero    Male                        -   
11                  Air-Walker    Male                     blue   
12                        Ajax    Male                    brown   
13                  Alan Scott    Male                     blue   
14                 Al
```

```python
bhbe = 
bhbe.sum()
```

Output:
```
93
```

```python
good = heroes[heroes['Alignment'] == 'good']
good_b = good[(good['Eye color'] == 'blue') & (
    (good['Hair color'].str.contains('Blond') | (good['Hair color'].str.contains('blond'))))]
p_good = len(good_b) / len(good)

p_all= len(good)/len(heroes)
p_all - p_good
```

Output:
```
0.5164751252527028
```

```python

```

Output:
```
496
```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
superheroes_fp = os.path.join('data', 'superheroes.csv')
heroes = pd.read_csv(superheroes_fp, index_col=0)
bhbe_out = bhbe_col(heroes)

obs_stat_out = observed_stat(heroes)

simulate_bhbe_out = simulate_bhbe_null(10)

pval_out = calc_pval()
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

### Finish Line

Before submitting your lab, make sure to run the doctests in the terminal with `python -m doctest lab.py`. If all of the tests in the notebook pass, but some fail when uploading to Gradescope, make sure that you've run the doctests in the terminal and they all pass.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!

q4 results: All test cases passed!

q5 results: All test cases passed!

q6 results: All test cases passed!

q7 results: All test cases passed!

q8 results: All test cases passed!
```

```python

```

Tests for q1 (points: None)
```
>>> isinstance(scores, pd.DataFrame)
True
```
```
>>> list(scores.columns)
['attempts', 'highest_score']
```
```
>>> isinstance(scores.index[0], int)
False
```
```
>>> isinstance(passfail, pd.DataFrame)
True
```
```
>>> len(passfail.columns)
3
```
```
>>> passfail.loc["Julia", "pass"] == 'Yes'
True
```
```
>>> isinstance(avscore, float)
True
```
```
>>> 91 < avscore < 92
True
```
```
>>> isinstance(highest, dict)
True
```
```
>>> len(next(iter(highest.items()))[1])
3
```
```
>>> isinstance(idxdup, int)
True
```
```
>>> 1 <= idxdup <= 6
True
```


Tests for q2 (points: None)
```
>>> ans =  trick_me()
>>> ans == 'A' or ans == 'B' or ans == "C"
True
```
```
>>> ans = reason_dup()
>>> ans == 'A' or ans == 'B' or ans == "C"
True
```
```
>>> isinstance(trick_ans, list)
True
```
```
>>> isinstance(trick_ans[1], str)
True
```
```
>>> ans = reason_bool()
>>> ans == 'A' or ans == 'B' or ans == "C" or ans =="D"
True
```


Tests for q3 (points: None)
```
>>> change(1.0) == 1.0
True
```
```
>>> change(np.NaN) == 'MISSING'
True
```
```
>>> nans = pd.DataFrame([[0,1,np.NaN], [np.NaN, np.NaN, np.NaN], [1, 2, 3]])
>>> A = correct_replacement(nans)
>>> (A.values == 'MISSING').sum() == 4
True
```


Tests for q4 (points: None)
```
>>> out_pop.index.tolist() == ['A', 'B', 'C', 'D']
True
```
```
>>> cols = ['num_nonnull', 'pct_nonnull', 'num_distinct', 'pct_distinct']
>>> out_pop.columns.tolist() == cols
True
```
```
>>> (out_pop['num_distinct'] <= 10).all()
True
```
```
>>> (out_pop['pct_nonnull'] == 1.0).all()
True
```
```
>>> common_out.index.tolist() == [0, 1, 2]
True
```
```
>>> common_out.columns.tolist() == ['A_values', 'A_counts', 'B_values', 'B_counts']
True
```
```
>>> common_out['A_values'].isin(range(10)).all()
True
```


Tests for q5 (points: None)
```
>>> isinstance(null_hypoth(), list)
True
```
```
>>> set(null_hypoth()).issubset({1,2,3,4})
True
```
```
>>> pd.Series(simulate_null()).isin([0,1]).all()
True
```
```
>>> 0 < estimate_p_val(1000) < 0.1
True
```


Tests for q6 (points: None)
```
>>> isinstance(super_out, list)
True
```
```
>>> len(super_out)
3
```
```
>>> all([isinstance(x, str) for x in super_out])
True
```


Tests for q7 (points: None)
```
>>> clean_out['Skin color'].isnull().any()
True
```
```
>>> clean_out['Weight'].isnull().any()
True
```
```
>>> stats_out[0] in ['Marvel Comics', 'DC Comics']
True
```
```
>>> isinstance(stats_out[1], int)
True
```
```
>>> isinstance(stats_out[2], str)
True
```
```
>>> stats_out[3] in ['good', 'bad']
True
```
```
>>> isinstance(stats_out[4], str)
True
```
```
>>> 0 <= stats_out[5] <= 1
True
```


Tests for q8 (points: None)
```
>>> isinstance(bhbe_out, pd.Series)
True
```
```
>>> bhbe_out.dtype == np.dtype('bool')
True
```
```
>>> bhbe_out.sum()
93
```
```
>>> 0.5 <= obs_stat_out <= 1.0
True
```
```
>>> isinstance(simulate_bhbe_out, pd.Series)
True
```
```
>>> simulate_bhbe_out.shape[0]
10
```
```
>>> ((0.45 <= simulate_bhbe_out) & (simulate_bhbe_out <= 1)).all()
True
```
```
>>> len(pval_out)
2
```
```
>>> 0 <= pval_out[0] <= 1
True
```
```
>>> pval_out[1] in ['Reject', 'Fail to reject']
True
```
